Heterogeneity of use, access and retention of insecticide-treated nets: implications for subnational tailoring to maximise malaria control

Author response:

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public review):

This paper aims to improve the accuracy of predictions of the impact of ITN strategies by developing a method to estimate duration of ITN access and use over time on a subnational scale from cross-sectional survey data and the numbers ITNs received annually. The subnational estimates are then input into a mathematical model to predict clinical cases under different ITN distribution strategies.

Strengths:

The approach is novel and addresses a useful and timely topic. It makes use of available routine data, and has considered all of the relevant components of ITN distributions.

The authors have made revisions, particularly to the methods, appendices and title - leaving the paper easier to follow, and with a clear, consistent aim. The assumptions are clearly stated.

Weaknesses:

The weaknesses are shared with other models of a similar complexity - it is not easy for a casual reader to fully understand the model or the implications of the assumptions which were required to be made. That routine data is used is good for availability, but data quality may be an issue in some places.

Reviewer #2 (Public review):

Summary:

The authors design a custom Bayesian model to estimate the probabilities of access, use and use given access of insecticide-treated nets in six African countries, providing sub-national estimates and inferring the average duration of ITN use and access. An individual-based model was employed to simulate malaria epidemics and estimate the effectiveness of different ITN distribution strategies. The study finds that the mean probability of use or access did not reach 80% (a universal coverage formerly targeted by WHO) for any of the regions even for biennial campaigns, demonstrates that switching from triennial to biennial distribution campaigns increases population use by 7.9%, and evaluates the impact of employing more efficient ITNs on P. falciparum prevalence.

Strengths:

The authors developed a data-driven model that accounts for data collection imperfections and sources of uncertainty while differentiating between ITN use and access. They developed a methodology to infer the timing of mass campaign from publicly available data instead of assuming fixed dates. The probability of use given access allows determining the regions where ITN distribution is least effective. This work can help better inform future interventions by identifying regions where increasing mass campaign frequency or employing better ITNs are most effective. Finally, in addition to insights on ITN access and use for the six countries analyzed, the paper contributes with a methodological framework that can likely be extended to other countries.

Weaknesses:

Since the models employed are rather complex, the methodology description may be hard to follow for some readers. In addition, the models assume many hypotheses, including exponential decay of ITN use/access and narrow prior distributions. It is worth noting that, in the revised version of the manuscript, the authors justified the choice of exponential decay and narrow prior distributions, and made a significant effort to clarify the methodology and the model equations.

Comments on revised version:

I appreciate the improvements made to the text. The methodology description is much clearer now. I have no further suggestions.

We thank the reviewers and editors for their constructive and insightful comments throughout the review process.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

P8 'Improving ITN use' L218

The numbers do not seem add up to me. "...increases across all settings of 14.5% (95% CrI:14.5, 14.6), from 41.7%% to 49.6%. Greater increases are predicted to be seen for ITN use with mean use across all settings increasing from 58.0% to 66.2%, an increase of 19.5% CrI (95% CrI:19.5, 19.6)."

Thank you for highlighting this. We have reviewed all reported results on mean use, access and use given access. The previous text reported a mixture of absolute and relative % changes, as well as a mixture of raw mean estimates across all regions and population-weighted means across regions. In the extract above we had inadvertently mixed different metrics. Given administrative-one regions can vary notably in population between different countries, we have ensured estimates are now consistently reported as population-weighted means, so that countries with finer-scaled administrative-one regions, such as Burkina Faso, do not artificially bias a raw mean estimate across all sub-national regions. We have also reported % changes as absolute percentage-point increases throughout, rather than relative ones to improve clarity.

Methods p18: There is notation in the text which does not seem to be explained. It is in the appendices, but the appendices should be optional extra information rather than essential for understanding.

We have reviewed the text in the main methods to check notation explanations. Following this, we have removed a use of subscript i, which is only used in the appendices to explicitly indicate region-specific parameters, and have clarified that lambda is a decay parameter.

There are assumptions made and these are clearly explained in the text. However, how much the highlighted results rest on the assumptions was not clear, and there was little on this in the discussion.

For example, it might seem disappointing that changing from triennial to biennial ITN campaigns would only lead to an increase from 41.7% to 49.6%. The most important assumptions driving this could be clearer. Additionally, after reading I was not sure what the likely consequences of the assumption that ITN are used continuously were.

We have added some additional text “to the discussion to clarify the modest predicted increase under biennial campaigns may, in part, be influenced by our assumed exponential loss function, and have highlighted that larger increases in mean use could plausibly be predicted under alternative ITN loss functions”. However, we have also commented that our mean use estimates are broadly in agreement with time series modelled estimates by Bertozzi-Villa et al. (2021) who utilised a sigmoidal/smooth-compact loss function.

In relation to the assumption of continuous use, we have added additional text in the ‘Historical use, access and retention times’ methods section to clarify that “if ITN use were systematically higher during high-transmission rainy seasons, our assumption of continuous use may underestimate the protective impact of ITNs during these periods”. As stated at the start of that paragraph, the data available from DHS surveys was too infrequent to investigate seasonal fluctuations.

P14 The text seems to imply that current transmission intensity is the only criterion for decisions about interventions. However, it is likely that the reasons for the current intensity, such as vectorial capacity, historical transmission and interventions should also play a role. The wording could reflect this.

We have added additional text to clarify that current transmission intensity should not be treated as the only criterion for deprioritisation decisions:

“However, current incidence should be considered alongside the factors that gave rise to that transmission intensity, with caution exercised when deprioritising mass campaigns in areas where historically higher transmission may currently be suppressed by high ITN access, high use given access, or other interventions.”

Minor points

There are several definite numbers in the first paragraph of the Introduction - these are estimates rather than the absolute truth, but the wording does not acknowledge that there is uncertainty.

We have made minor edits to clarify that these values are estimates rather than exact quantities. Measures of uncertainty, such as credible intervals were not always possible to source; for example, some of these are median estimates inferred from figures in Bertozzi-Villa et al. (2021).

L634 typo - logisitic

Now corrected.

L1731 typo https://https://

Now corrected.

L881 "access at random" - perhaps not the easiest for non-modelers

We have re-written this to clarify “when ITNs in a household can provide access to more individuals than the number of users, access is assigned at random to non-users within each household under our framework”.

Appendix 1, table 1: Using alpha for both age and also overdispersion on use or access is of course valid, but I found it a little confusing.

To avoid confusion, we have added the following clarification in brackets:

“Meanwhile, the overdispersion parameter, $\alpha_i^0$ (unrelated to the notation for ITN age), controls the variability of the probability of individual access around the mean”

I suspect that the model was actually fitted in Stan via the R interface rstan (L589, L1151 and elsewhere).

We have now clarified this throughout.

Changes from the earlier version to the revised version

Abstract

Earlier version: Insecticide-treated nets (ITNs) are the most impactful and cost-effective control tool against malaria. ITNs are primarily distributed through triennial mass campaigns across Africa, though overall ITN use remains modest in many areas as most ITNs do not last three years. In times of funding constraints and a lack of economic alternative antimalarial interventions it is unclear whether disease control could be best improved by distributing more effective ITNs (e.g. dual active-ingredient ITNs) and/or deploying nets more frequently. There are increased calls to improve allocation of resources through sub-national tailoring of interventions, though benefits will depend on how long people use ITNs and how this varies between subnational regions. However, subnational variation in ITN retention and the duration that ITNs remain in use have not previously been quantified. Here we estimate subnational differences in ITN use, access and retention for six countries, Burkina Faso, Ghana, Malawi, Mali, Mozambique and Senegal. These estimates are used to calibrate a Plasmodium falciparum transmission dynamics model to generate sub-national estimates of ITN use and cases averted under different ITN distribution strategies. On average, people use their ITNs for 21 months, though this varies substantially between subnational regions from 12 to 38 months. Shifting from triennial to biennial campaigns is predicted to lead to mean population use across all regions increasing from 41.7% to 49.6%. No regions of the 146 investigated were estimated to maintain use over 80% even under biennial distribution, though switching to dual active-ingredient ITNs would likely avert more cases under present distribution frequencies. Our results highlight that although transmission intensity remains an important factor for subnational tailoring of malaria control interventions, other factors, such as ITN use given access, meaningfully influence optimal deployment strategies. The framework highlights how routinely collected data can aid policymakers in tailoring disease control programmes at subnational levels.

Revised version: Insecticide-treated nets (ITNs) are the most impactful and cost-effective control tool against malaria. ITNs are primarily distributed through triennial mass campaigns across Africa, though overall ITN use remains modest in many areas as most ITNs do not last three years. In times of funding constraints and a lack of economic alternative antimalarial interventions it is unclear whether disease control could be best improved by distributing more effective ITNs (e.g. dual active-ingredient ITNs) and/or deploying nets more frequently. There are increased calls to improve allocation of resources through sub-national tailoring of interventions, though benefits will depend on how long people use ITNs and how this varies between subnational regions. However, subnational variation in ITN retention and the duration that ITNs remain in use have not previously been quantified. Here we estimate subnational differences in ITN use, access and retention for six countries, Burkina Faso, Ghana, Malawi, Mali, Mozambique and Senegal. These estimates are used to calibrate a Plasmodium falciparum transmission dynamics model to generate sub-national estimates of ITN use and cases averted under different ITN distribution strategies. On average, people use their ITNs for 21 months, though this varies substantially between subnational regions from 12 to 38 months. Shifting from triennial to biennial campaigns is predicted to lead to mean population-weighted use across all regions increasing from 45.4% to 53.9%. No regions of the 146 investigated were estimated to maintain use over 80% even under biennial distribution, though switching to dual active-ingredient ITNs would likely avert more cases under present distribution frequencies. Our results highlight that although transmission intensity remains an important factor for subnational tailoring of malaria control interventions, other factors, such as ITN use given access, meaningfully influence optimal deployment strategies. The framework highlights how routinely collected data can aid policymakers in tailoring disease control programmes at subnational levels.

Introduction

Earlier version: The widespread use of insecticide-treated nets (ITNs) has prevented more malaria than any other intervention. For populations at risk of malaria, national malaria programmes (NMPs) have historically aimed to achieve ‘universal coverage’ of ITNs, which is defined as a measure of both universal access to, and use of, ITNs (World Health Organization, 2014). The 2005-2015 Roll Back Malaria Partnership (2005) Global Strategic Plan first proposed universal coverage targets of 80% of people at risk of malaria by 2010 (Roll Back Malaria Partnership, 2008; Smith et al., 2009). The World Health Organization (2024b) Guidelines for Malaria recommended achieving universal coverage by implementing mass distribution campaigns, and through the use of complimentary continuous distribution channels; these are typically conducted via antenatal care (ANC), the Expanded Programme for Immunization (EPI) and through school-based distribution (World Health Organization, 2024b; Koenker et al., 2022). Mass campaigns have historically been recommended to be conducted every three years with a defined number of nets distributed per person (World Health Organization, 2024b; Koenker et al., 2022). These policies have increased estimated ITN use across sub-Saharan Africa from 2% to 49% between 2000 to 2022 (Bertozzi-Villa et al., 2021). Access to a net (assuming an ITN could be used by up to two people) has also gone up from 3% to 56% over the same period. While this falls short of the original recommendation, ITNs have been instrumental in reducing the global malaria burden, having averted an estimated 450 million cases between 2000 and 2015 (Bhatt et al., 2015), or two-thirds of all cases averted through preventative interventions over the same period. Nevertheless, the disparity between recent estimates of ITN use and the original 80% targets are stark. High population use may be unachievable with the status quo of triennial campaigns, since it is thought people do not use their nets for three years, with previous median estimates of national median ITN retention times being lower in 87.5% of sub-Saharan African countries (Bertozzi-Villa et al., 2021). There is also likely to be substantial subnational variability due to environmental and behavioural differences, though evidence is lacking due to the sporadic nature of net surveillance activities. Since 2019, WHO good practice statements have shifted from advocating for ‘universal’ to ‘optimal’ coverage guided by “an explicit prioritization process guiding resource allocation decisions” (World Health Organization, 2024b). However, it is unclear what the optimum coverage is, how it can be estimated, and how this might vary with net retention time.

Revised version: The widespread use of insecticide-treated nets (ITNs) has prevented more malaria than any other intervention. For populations at risk of malaria, national malaria programmes (NMPs) have historically aimed to achieve ‘universal coverage’ of ITNs, which is defined as a measure of both universal access to, and use of, ITNs (World Health Organization, 2014). The 2005-2015 Roll Back Malaria Partnership (2005) Global Strategic Plan first proposed universal coverage targets of 80% of people at risk of malaria by 2010 (Roll Back Malaria Partnership, 2008; Smith et al., 2009). The World Health Organization (2024b) Guidelines for Malaria recommended achieving universal coverage by implementing mass distribution campaigns, and through the use of complimentary continuous distribution channels; these are typically conducted via antenatal care (ANC), the Expanded Programme for Immunization (EPI) and through school-based distribution (World Health Organization, 2024b; Koenker et al., 2022). Mass campaigns have historically been recommended to be conducted every three years with a defined number of nets distributed per person (World Health Organization, 2024b; Koenker et al., 2022). These policies are estimated to have increased ITN use across sub-Saharan Africa from median point estimates of approximately 2% to 49% between 2000 to 2021 (Bertozzi-Villa et al., 2021). Access to a net (assuming an ITN could be used by up to two people) has also similarly been estimated to have increased from median estimates of 3% to 56% over the same period (Bertozzi-Villa et al., 2021). While this falls short of the original recommendation, ITNs have been instrumental in reducing the global malaria burden, having averted an estimated 450 million cases between 2000 and 2015 (Bhatt et al., 2015), or two-thirds of all cases averted through preventative interventions over the same period. Nevertheless, the disparity between recent estimates of ITN use and the original 80% targets are stark. High population use may be unachievable with the status quo of triennial campaigns, since it is thought people do not use their nets for three years, with previous median estimates of national median ITN retention times being lower in 87.5% of sub-Saharan African countries (Bertozzi-Villa et al., 2021). There is also likely to be substantial subnational variability due to environmental and behavioural differences, though evidence is lacking due to the sporadic nature of net surveillance activities. Since 2019, WHO good practice statements have shifted from advocating for ‘universal’ to ‘optimal’ coverage guided by “an explicit prioritization process guiding resource allocation decisions” (World Health Organization, 2024b). However, it is unclear what the optimum coverage is, how it can be estimated, and how this might vary with net retention time.

The epidemiological impact of mass distributions of ITNs will depend on how long individuals retain and sleep under nets and how effective ITNs are at killing mosquitoes and inhibiting blood-feeding. The original rationale for triennial mass campaigns is unclear and is likely based on the durability of first-generation ITNs treated solely with pyrethroid-class insecticides (World Health Organization, 2005, 2013). The real-world effectiveness of pyrethroid-based ITNs is under threat from increasingly high levels of pyrethroid resistance in anopheline mosquitoes across sub-Saharan Africa, which is highly heterogenous over national and sub-national scales (Coleman et al., 2017; Moyes et al., 2020; Hancock et al., 2020, 2022). Experimental hut trials have indicated increased pyrethroid resistance in the vector population not only reduces the initial mortality conferred by new ITNs, but also increases the rate they lose their insecticidal properties following distribution (Nash et al., 2021; Sherrard-Smith et al., 2022; Churcher et al., 2024). WHO recommendations suggest that more regular mass campaigns or altering the distribution ratio might be appropriate given local data-driven evidence (World Health Organization, 2024b; Koenker et al., 2022). Nevertheless, international procurers of ITNs, such as the Global Fund, have not historically advised altering the regularity of ITN mass campaigns, nor is this considered in sub-national tailoring exercises, despite this potentially being more cost-effective than layering multiple interventions in the same region.

Earlier version: The World Health Organization has recently expanded its guidance on different types of ITN, with three classes now recommended in areas of pyrethroid resistance. WHO recommendations cascade in the following order: i) pyrethroid-chlorfenapyr nets are strongly recommended for use in place of the standard pyrethroid-only nets; ii) pyrethroid-piperonyl butoxide (PBO) treated nets are conditionally recommended in favour of the standard pyrethroid-only nets, and iii) pyrethroid-pyriproxifen (PPF) treated nets are conditionally recommended in favour of pyrethroid-only nets, but are recommended not to be used in place of pyrethroid-PBO nets (World Health Organization, 2024b). These recommendations are based on results from cluster randomised controlled trials showing that pyrethroid-PBO (Protopopoff et al., 2018; Staedke et al., 2020) and pyrethroid-chlorfenapyr (Accrombessi et al., 2023; Churcher et al., 2024) products can both mitigate pyrethroid-resistance associated decreases in ITN effectiveness, while superiority trials for pyrethroid-PPF nets showed mixed results (Accrombessi et al., 2023). The durability of different classes of nets is also thought to vary substantially, further exacerbating the need to consider the optimum regularity of mass campaigns.

Revised version: The World Health Organization has recently expanded its guidance on different types of ITN, with three classes now recommended in areas of pyrethroid resistance. WHO recommendations cascade in the following order: i) pyrethroid-chlorfenapyr nets are strongly recommended for use in place of the standard pyrethroid-only nets; ii) pyrethroid-piperonyl butoxide (PBO) treated nets are conditionally recommended in favour of the standard pyrethroid-only nets, and iii) pyrethroid-pyriproxifen (PPF) treated nets are conditionally recommended in favour of pyrethroid-only nets, but are recommended not to be used in place of pyrethroid-PBO nets (World Health Organization, 2024b). These recommendations are based on results from cluster randomised controlled trials showing that pyrethroid-PBO (Protopopoff et al., 2018; Staedke et al., 2020) and pyrethroidchlorfenapyr (Accrombessi et al., 2023; Churcher et al., 2024) products can both mitigate pyrethroid-resistance associated decreases in ITN effectiveness, while superiority trials for pyrethroid-PPF nets showed mixed results (Accrombessi et al., 2023). The durability of different classes of nets is also thought to vary substantially, further exacerbating the need to consider the optimum regularity of mass campaigns.

NMPs therefore face two key challenges for planning future ITN interventions. Firstly, to increase ITN use, more frequent mass campaigns may be necessary, given ITNs are typically retained for shorter durations than the status quo mass campaign interval of three years. Secondly, in Africa, where pyrethroid resistance is widespread, newer, yet costlier ITN products are advisable. Addressing either of these challenges would require increases in funding, yet budgets are increasingly under threat and there are increasing calls to tailor ITN distribution within countries to prevent waste and maximise benefit. This is recommended by the WHO (World Health Organization, 2024a) but has not yet been implemented at scale. The impact of both these policy choices may not only depend on how access to ITNs, pyrethroid resistance and transmission intensity vary between and within countries, but may also depend on how likely individuals with access to ITNs are to use them, in addition to the duration that ITNs continue to be used. It is unclear which of these factors are likely to be most associated with improvements in health outcomes through more regular mass campaigns, or more effective ITNs. Moreover, it is unknown if net retention varies systematically within a country or if ITNs remain in use for shorter durations than they are retained for.

To investigate optimal ITN deployment, we estimate the mean duration of ITN access and use from Demographic and Health Survey (DHS) data for different subnational regions in six African countries (Burkina Faso, Ghana, Malawi, Mali, Mozambique and Senegal). Here, (subnational) regions are defined as the first administrative unit below the country level and are further divided into rural and urban areas to align with DHS stratification. A bespoke mechanistic statistical model is fitted to subnational DHS data from 2005 to 2024. It allows regions with sporadic DHS surveys to be directly compared and generates estimates of ITN use, access and use given access over time, in addition to the contribution towards these by ITNs distributed through continuous channels (ITNs obtained through the private sector, any school- or community-based distribution programme, or through routine ANC and EPI channels). Projections are made for how these metrics change from shifting from triennial to biennial mass campaigns. Using these subnational estimates we calibrated individual-based models of Plasmodium falciparum malaria transmission to generate subnational projections of clinical cases averted over 2025 to 2030 under different ITN distribution strategies. Our findings should enable assessment of the epidemiological benefit of different ITN strategies and provide data-driven evidence to improve ITN priorisation guidance.

Results

Retention time and duration of use

The model estimates people keep and use ITNs for considerably less than three years. Across all regions in all six countries studied, the population-weighted mean duration of ITN access, herein referred to as ITN retention time, was 26.8 months (95% CrI: 26.7, 26.9). Retention times were estimated to vary notably between and within countries (figure 1), with median estimates ranging from 14.8 months (95% CrI: 14.5, 15.1) in rural Tete, Mozambique, to 41.5 months (95% CrI: 40.7, 42.2) in rural Ségou, Mali, with only 23 of 146 (15.8%) regions estimated to have a mean retention time greater than three years. The mean durations of ITN use similarly varied notably both subnationally and between countries and were estimated to be consistently lower than mean retention times (figure 1), with a population-weighted mean across all regions of 21.0 months (95% CrI: 20.9, 21.0). The mean duration that ITNs remain in use was also estimated to be shortest in rural Tete with a median estimate of 12.0 months (95% CrI: 11.8, 12.3). Meanwhile, the longest mean duration of use was in urban Est, Burkina Faso (figure 5 – figure supplement 1.A) at 37.8 months (95% CrI: 34.9, 40.6), which was the only region (0.7%) with a median estimate greater than three years. Caution is needed when comparing retention times between countries, due to differences in the timings of surveys relative to mass campaigns, but these are likely to be minimised over subnational scales within a country. While no systematic differences in the durations of use or access were estimated between urban and rural regions in Senegal, differences were observed in some other countries, such as in Ghana, where values were estimated to be broadly lower in urban settings (figure 5 – figure supplement 2.A). For most regions, the gap between use and access was estimated to widen over time following mass campaigns, indicating people stop using ITNs substantially before they stop owning them, so that use given access is not constant but typically wanes as ITNs age (figure 2.A).

Population-weighted distribution of expected ITN retention times and duration of use by country.

Median estimates of the mean ITN retention time and duration of use for each subnational region were weighted by population to produce country-level distributions. These histograms represent probability distributions of retention time and duration of use for an individual randomly selected from each country, assuming model-derived median estimates for all subnational regions. Subnational variability in retention times and the duration of use reflects differences in model-derived estimates between regions, since the model assumes these values are constant within regions. Vertical dashed lines indicate the overall population-weighted mean (black) and the population-weighted mean for each country (coloured).

ITN use, access and Pf PR over time in rural Burkina Faso.

Earlier version: A: Survey estimates of ITN use (probability someone slept under an ITN the previous night, points sized proportionally to sample size) and model mean estimates of the probability of use over time (solid lines) for rural portions of subnational regions in Burkina Faso. Mean estimates are shown for the probability of access (dotted lines) and the probability of using an ITN sourced from a continuous channel (dashed lines).The model accounts for differing probabilities of use between individuals within a region; this results in notable uncertainty around the probability an individual selected at random uses a ITN. This uncertainty is illustrated in the shaded region, with the 50%, 80% and 95% credible intervals as indicated by progressively lighter shaded regions. B: Pf PR2−10 estimates from the transmission dynamics model (darker lines) were fitted to annual estimates of Pf PR2−10 from the Malaria Atlas Project (2024) (hollow points); model estimates of Pf PR6−59mo are also shown (lighter lines), in addition to observed DHS Pf PR6−59mo and associated 95% credible intervals due to measurement uncertainty (solid points and vertical lines). Figure 2—figure supplement 1. Use, access and Pf PR over time in rural Senegal. Figure 2—figure supplement 2. Use of ITNs from continuous channels.

Revised version: A: Survey estimates of ITN use (probability someone slept under an ITN the previous night, points sized proportionally to sample size) and model mean estimates of the probability of use over time (solid lines) for rural portions of subnational regions in Burkina Faso. Mean estimates are shown for the probability of access (dotted lines) and the probability of using an ITN sourced from a continuous channel (dashed lines).The model accounts for differing probabilities of use between individuals within a region; this results in notable uncertainty around the probability an individual selected at random uses a ITN. This uncertainty is illustrated in the shaded region, with the 50%, 80% and 95% credible intervals as indicated by progressively lighter shaded regions. B: Pf PR2−10 estimates from the transmission dynamics model (darker lines) were fitted to annual estimates of Pf PR2−10 from the Malaria Atlas Project (2024) (hollow points); model estimates of Pf PR6−59mo are also shown (lighter lines), in addition to observed DHS Pf PR6−59mo and associated 95% credible intervals due to measurement uncertainty (solid points and vertical lines).

Universal coverage: was it achievable under triennial mass campaigns?

Access and use of ITNs are estimated to have varied substantially over time, both within and between countries (figure 2.A and 2 - figure supplement 1.A). Inferring changes over time from observational data alone can give misleading results. For example, in rural Centre-Nord in Burkina Faso, DHS survey data suggests a sustained decrease in use between 2014 and 2022 (figure 2.A); however, the model which considered the timings of mass campaigns and routine ITN distribution estimates that overall ITN use immediately following mass campaigns increased over this period. The observed decrease in use and access across many regions in Burkina Faso may therefore be a by-product of DHS surveys being conducted at progressively later dates relative to the most recent campaign; this does not necessarily indicate an underlying trend in decreasing use or access over longer timescales. This highlights the potential limitations of making direct comparisons between different surveys, which may be conducted at different intervals since the preceding campaign (figure 2.A). This waxing and waning dynamics in use over time is also observed in P. falciparum prevalence (figure 2.B). In many countries peak ITN use and access following mass campaigns has increased over time; for example in rural Thiès in Senegal, the overall use immediately following the first 2009 mass campaign was estimated to reach 18%, yet peak levels of use above 80% were estimated to be achieved following all mass campaigns since 2016 (figure 2 - figure supplement 1.A). Similar trends can be seen in the context of access (dotted lines, figure 2.A and figure 2 - figure supplement 1.A).

Earlier version: Mean ITN use over the three years following a campaign is estimated to be modest, with a population-weighted mean of 41.7% (95% CrI: 39.6, 43.7) across all countries studied. Overall access to an ITN was estimated to be greater over the three year period, with 58.0% (95% CrI: 57.9, 58.1) of people having access to an ITN, on average. The population-weighted mean use given access across all regions was estimated to be 71.9% (95% CrI: 68.2, 75.3). While no regions were estimated to have a mean use or access greater than 80% over either the three-year or two-year periods following a mass campaign, in the context of use given access, 46% of subnational regions had median estimates greater than 80% over a three-year period following a mass campaign, increasing to 70% of regions when only the first two years were considered (figure 3 and figure 3 – figure supplements 1–3). Nevertheless, some regional differences are observable, with southern Malawi having a low overall ITN use despite the region having a similar level of access to the rest of the country. Across all regions, mass campaigns are estimated to increase ITN use in their immediate aftermath to greater than 80% in 76.8% of regions (bottom row, figure 3 and figure 3 – figure supplements 1–3), with considerable variability within and between countries. For example, for rural areas in Burkina Faso, the northern Sahel region is predicted to never achieve ITN use above 80%, whereas use is predicted to be maintained above 80% for 6 months following a mass campaign in the south-western region of Cascades.

Revised version: Mean ITN use over the three years following a campaign is estimated to be modest, with a population-weighted mean of 45.4% (95% CrI: 45.2, 45.5) across all regions studied. Overall access to an ITN was estimated to be greater over the three year period, with 58.1% (95% CrI: 58.0, 58.2) of people having access to an ITN, on average. The population-weighted mean use given access across all regions was estimated to be 75.7% (95% CrI: 75.5, 76.0). While no regions were estimated to have a mean use or access greater than 80% over either the three-year or two-year periods following a mass campaign, in the context of use given access, 46% of subnational regions had median estimates greater than 80% over a three-year period following a mass campaign, increasing to 70% of regions when only the first two years were considered (figure 3 and figure 3 – figure supplements 1–3). Nevertheless, some regional differences are observable, with southern Malawi having a low overall ITN use despite the region having a similar level of access to the rest of the country. Across all regions, mass campaigns are estimated to increase ITN use in their immediate aftermath to greater than 80% in 76.8% of regions (bottom row, figure 3 and figure 3 – figure supplements 1–3), with considerable variability within and between countries. For example, for rural areas in Burkina Faso, the northern Sahel region is predicted to never achieve ITN use above 80%, whereas use is predicted to be maintained above 80% for 6 months following a mass campaign in the south-western region of Cascades.

Mean use and access with three-year campaigns in subnational rural areas.

Earlier version: Central estimates of mean overall proportion of people using an ITN the previous night, access to an ITN, and use given access for three-year mass campaign intervals are shown in the top three rows. The number of months following a mass campaign where overall ITN use exceeds 80% are shown in the bottom row. Figure 3—figure supplement 1. Mean use and access with three-year campaigns in subnational urban areas. Figure 3—figure supplement 2. Mean rural use and access with two-year campaigns. Figure 3—figure supplement 3. Mean urban use and access with two-year campaigns.

Revised version: Central estimates of mean overall proportion of people using an ITN the previous night, access to an ITN, and use given access for three-year mass campaign intervals are shown in the top three rows. The number of months following a mass campaign where overall ITN use exceeds 80% are shown in the bottom row.

Earlier version: The contribution of ITNs from continuous channels also differed between and within countries; for example, in rural regions of Thiès (figure 2 - figure supplement 1) the use of ITNs from continuous channels was estimated to be 15.8% (95% CrI: 14.7, 17.0) at the end of 2024, in comparison to 4.1% (95% CrI: 3.8,4.5) in Sédhiou. Differences in the use of continuously-distributed ITNs were also notable between countries, with 19.8% (95% CI: 17.4, 22.2) of person-nights of ITN use (one person using an ITN for one night) across Burkina Faso estimated to be attributable to continuously-distributed ITNs over a three-year period, compared to Malawi at 61.2% (95% CI: 58.3, 64.1) (figure 2 - figure supplement 2.B).

Revised version: The contribution of ITNs from continuous channels also differed between and within countries; for example, in rural regions of Thiès (figure 2 figure supplement 1) the use of ITNs from continuous channels was estimated to be 15.8% (95% CrI: 14.7, 17.0) at the end of 2024, in comparison to 4.1% (95% CrI: 3.8,4.5) in Sédhiou. Differences in the use of continuously-distributed ITNs were also notable between countries, with 19.8% (95% CI: 17.4, 22.2) of person-nights of ITN use (one person using an ITN for one night) across Burkina Faso estimated to be attributable to continuouslydistributed ITNs over a three-year period, compared to Malawi at 61.2% (95% CI: 58.3, 64.1) (figure 2 - figure supplement 2.B).

There is considerable variability in observed mean access and use within a region, with notable uncertainty between communities, as indicated by the high uncertainty around the probability an individual randomly selected from the region will use an ITN (figure 1.A). The effect of this is illustrated by rural Thiès and Ziguinchor, which both have a mean overall ITN use of 58% when averaged over three years following the last mass campaign. Though the average is the same across both regions, more people in rural areas of Ziguinchor either have a very low or very high probability of using an ITN, indicating greater inequity in ITN use in rural areas of Ziguinchor than in Thiès (figure 4). A similar trend is also apparent in the context of access and in other settings (figure 4 - figure supplement 1 and 2).

Equity of use and access in rural Thiès and Ziguinchor, Senegal.

Earlier version: Solid vertical lines denote overall ITN use and access averaged over three years following the last mass campaign. Coloured bars indicate the proportion of the population with different probabilities of using or having access to an ITN. For example, 10% of the population have a probability of access between 70% and 80% in rural Ziguinchor, in comparison to 20% in rural Thiès. Vertical dashed lines denote 50% credible intervals for the probability of an individual uses or has access to an ITN, indicating that half of the population are expected to have a probability of use or access within this range. Figure 4—figure supplement 1. Subnational equity of use in Senegal. Figure 4—figure supplement 2. Subnational equity of access in Senegal.

Revised version: Solid vertical lines denote overall ITN use and access averaged over three years following the last mass campaign. Coloured bars indicate the proportion of the population with different probabilities of using or having access to an ITN. For example, 10% of the population have a probability of access between 70% and 80% in rural Ziguinchor, in comparison to 20% in rural Thiès. Vertical dashed lines denote 50% credible intervals for the probability of an individual uses or has access to an ITN, indicating that half of the population are expected to have a probability of use or access within this range.

Improving ITN use

The framework can identify regions of low ITN use and make suggestions about how this could be improved (figure 5.A,B). ITNs will be most effective in areas where people have access to nets for a long period of time and are more likely to use them (quadrant 1, figure 5.B). In areas where people keep ITNs but have low use given access (quadrant 2, figure 5.B) NMPs may wish to consider social and behaviour change (SBC) interventions such as ‘hang-up’ campaigns. Conversely, in areas with lower ITN retention times but high use given access (quadrant 4, figure 5.B) NMPs may wish to consider more regular mass campaigns to improve equality of access between different regions. Both of these policy changes may be warranted in regions with both low retention and use given access to improve overall ITN use (quadrant 3, figure 5.B). Results varied by country (figure 5 – figure supplements 1-5).

Retention time, use given access and changes in cases for Mali.

Earlier version: All coloured points indicate median estimates and associated 95% credible intervals for subnational regions in Mali. A: Mean durations of ITN use (hollow) and retention (filled) are shown in addition to national median estimates (black solid and dashed lines, respectively). B: Mean ITN retention and use given access under a triennial campaign strategy are shown, with black lines indicating national median estimates. C: Points indicate mean annual clinical cases under a triennial pyrethroid–PBO campaign strategy, and the projected change in cases for alternative intervention strategies with pyrethroid-only (C.i–iii) pyrethroid-PBO (C.iv–vi) or pyrethroid–chlorfenapyr (C.vii–ix) ITNs, with continuous-only distribution (C.i, iv, vii), or in conjunction with triennial (C.ii, v, viii) or biennial (C.iii, vi, ix) campaigns. The change in clinical cases is equal to zero for equivalent comparator and intervention strategies (C.v). Labeled diagonal reference lines with positive gradients indicate percentage increases in clinical cases relative to mean annual estimates under the comparator strategy; lines of the same style with negative gradients indicate equivalent percentage decreases. Urban areas were estimated to broadly have lower use given access than in rural settings, while the capital city, Bamako, was estimated to have the lowest ITN retention time of all regions. Figure 5—figure supplement 1. Retention time, use given access and changes in cases for Burkina Faso. Figure 5—figure supplement 2. Retention time, use given access, and changes in cases for Ghana. Figure 5—figure supplement 3. Retention time and use given access for Malawi. Figure 5—figure supplement 4. Retention time, use given access and changes in cases for Mozambique. Figure 5—figure supplement 5. Retention time, use given access and changes in cases for Senegal.

Revised version: All coloured points indicate median estimates and associated 95% credible intervals for subnational regions in Mali. A: Mean durations of ITN use (hollow) and retention (filled) are shown in addition to national median estimates (black solid and dashed lines, respectively). B: Mean ITN retention and use given access under a triennial campaign strategy are shown, with black lines indicating national median estimates. C: Points indicate mean annual clinical cases under a triennial pyrethroid–PBO campaign strategy, and the projected change in cases for alternative intervention strategies with pyrethroid-only (C.i–iii) pyrethroid-PBO (C.iv–vi) or pyrethroid–chlorfenapyr (C.vii–ix) ITNs, with continuous-only distribution (C.i, iv, vii), or in conjunction with triennial (C.ii, v, viii) or biennial (C.iii, vi, ix) campaigns. The change in clinical cases is equal to zero for equivalent comparator and intervention strategies (C.v). Labeled diagonal reference lines with positive gradients indicate percentage increases in clinical cases relative to mean annual estimates under the comparator strategy; lines of the same style with negative gradients indicate equivalent percentage decreases. Urban areas were estimated to broadly have lower use given access than in rural settings, while the capital city, Bamako, was estimated to have the lowest ITN retention time of all regions.

Earlier version: The statistical model can be used to predict the increase in use caused by increasing the frequency of mass campaigns from three to two years, assuming the contribution of continuous distribution channels and the decay in use given access following mass campaigns remains the same. Increases in access are estimated to be more modest, with relative population-weighted mean increases across all settings of 14.5% (95% CrI: 14.5, 14.6), from 41.7% to 49.6%. Greater increases are predicted to be seen for ITN use, with mean use across all settings increasing from 58.0% to 66.2%, an increase of 19.5% (95% CrI: 19.5, 19.6). Subnational variation is also predicted in changes of access, use and use given access following adoption of biennial mass campaigns (figure 3 and corresponding figure supplements 1-3).

Revised version: The statistical model can be used to predict the increase in use caused by increasing the frequency of mass campaigns from three to two years, assuming the contribution of continuous distribution channels and the decay in use given access following mass campaigns remains the same. Increases in access are estimated to be modest, with percentage-point population-weighted mean increases across all settings of 8.15% (95% CrI: 8.13, 8.17), from 58.1% (95% CrI: 58.0, 58.2) to 66.3 (95% CrI: 66.2, 66.4)%. Greater increases are predicted to be seen for ITN use, with the population-weighted mean use across all settings increasing from 45.4% (95% CrI: 45.3, 45.5) to 53.9% (95% CrI: 53.7, 54.1), an increase of 8.50% (95% CrI: 8.47, 8.54). An increase in use given access is also projected, with the population-weighted mean increasing by 4.13% (95% CrI: 4.07, 4.20) in absolute terms from 75.7% (95% CrI: 75.5, 76.0) to 79.9% (95% CrI: 79.6, 80.2). Subnational variation is also predicted in changes of access, use and use given access following adoption of biennial mass campaigns (figure 3 and corresponding figure supplements 1-3).

Improving ITN effectiveness

Policy-makers may wish to optimise ITN distribution to maximise epidemiological impact. We generate estimates of clinical cases under different distribution strategies by simulating halting mass campaigns or changing the campaign interval from three to two years. Fifty eight percent of ITNs distributed in sub-Saharan Africa in 2023 were reported to have been pyrethroid-PBO nets (World Health Organization, 2024c); we therefore estimate changes in clinical cases for alternative strategies compared to the current triennial pyrethroid-PBO campaigns with supplementary continuous distribution, with triennial pyrethroid-chlorfenapyr strategies as a comparator also presented (figure 6 - figure supplement 1).

Added in the revised version: Change in cases vs. cases with triennial pyrethroid-PBO distribution.

Added in the revised version: Points represent urban and rural areas within each subnational region, are colored by country, and are sized proportionally to the mean annual number of ITNs distributed under each strategy. Each point shows median estimates of the change in clinical cases following a switch from triennial pyrethroid-PBO distribution to alternative strategies against the projected clinical cases under the triennial pyrethroid-PBO strategy. Concurrent distribution of continuous ITNs of the same class used in mass campaigns in assumed throughout. Vertical and horizontal lines indicate 95% credible intervals. Labeled diagonal reference lines with positive gradients indicate percentage increases in clinical cases relative to mean annual estimates under the comparator strategy; corresponding lines with negative gradients represent equivalent percentage decreases and share the same line style.

The number of clinical cases averted by the different strategies varies substantially, both within (figure 5.C and figure 5 - figure supplement 1-5, panel C) and between countries (figure 6). For example, in Mali switching from a triennial to biennial pyrethroid-PBO mass campaign is estimated to avert fewer than 50 additional cases per 1,000 people in Gao, but over 100 in Sikasso (figure 5.C.vi). The benefit of the different strategies is projected to differ in each setting, though broad trends are observable. Generally, fewer cases are projected to be averted under triennial pyrethroid-chlofenapyr campaigns than with biennial pyrethroid-PBO distribution, with 74.6% of regions having fewer clinical cases predicted when you increase the campaign frequency than when the net class is changed (figure 6 - figure supplement 2.F vs. H and regions with negative y-axis values in figure 6 - figure supplement 1.F). Triennial distribution of pyrethroid-chlorfenapyr ITNs was predicted to avert more cases than biennial distribution of pyrethroid-only ITNs for 94.3% of regions (figure 6 - figure supplement 2.H vs. C and regions with positive y-axis values in figure 6 - figure supplement 1.C). This was more consistent for policy decision-making than the differences between triennial pyrethroid-PBO distribution and biennial pyrethroid-only distribution, with distributing pyrethroid-PBO ITNs averting more cases in 65.6% of regions (figure 6 - figure supplement 2.E vs. C and regions with positive y-axis values in figure 6.C).

Removed in the revised version: Change in cases vs. cases with triennial pyrethroid-PBO distribution.

Removed in the revised version: Points represent urban and rural areas within each subnational region, are colored by country, and are sized proportionally to the mean annual number of ITNs distributed under each strategy. Each point shows median estimates of the change in clinical cases following a switch from triennial pyrethroid-PBO distribution to alternative strategies against the projected clinical cases under the triennial pyrethroid-PBO strategy. Concurrent distribution of continuous ITNs of the same class used in mass campaigns in assumed throughout. Vertical and horizontal lines indicate 95% credible intervals. Labeled diagonal reference lines with positive gradients indicate percentage increases in clinical cases relative to mean annual estimates under the comparator strategy; corresponding lines with negative gradients represent equivalent percentage decreases and share the same line style. Figure 6—figure supplement 1. Change in cases vs. cases with triennial pyrethroid-chlorfenapyr distribution. Figure 6—figure supplement 2. Cases averted under different ITN distribution strategies. Figure 6—figure supplement 3. Change in cases vs Pf PR with triennial pyrethroid-PBO distribution. Figure 6—figure supplement 4. Change in cases vs pyrethroid resistance. Figure 6—figure supplement 5. Change in cases vs duration of use. Figure 6—figure supplement 6. Change in cases vs retention time. Figure 6—figure supplement 7. Change in cases vs use. Figure 6—figure supplement 8. Change in cases vs access. Figure 6—figure supplement 9. Change in cases vs use given access.

It is unlikely that bespoke simulation modeling will be available for all policy decisions, but the models can indicate trends that can support sub-national tailoring of ITNs. The models indicate current clinical incidence is the best predictor of the number of cases averted for the different strategies. For example, for both pyrethroid-PBO (figure 6.F) and pyrethroid-chlofenapyr ITNs (figure 6 - figure supplement 1.I), the projected additional cases averted by switching to biennial distribution is estimated to gradually increase as clinical cases increase up to approximately 300 mean annual clinical cases per 1,000 people, though above this point trends are less obvious. Similar trends are seen in countries with a broad range of transmission intensities, such as Mali (figure 5.C.vi) and Mozambique (figure 5 - figure supplement 4.B.vi), though it is less clear in countries where are no low transmission regions, such as in Burkina Faso (for example, figure 5 - figure supplement 1.C.vi). Similar trends can be observed by using all-age malaria prevalence as an indicator to select regions which will have the greatest benefit of greater investment in ITNs (figure 6 - figure supplement 3). Selecting areas to target using the level of pyrethroid resistance is less reliable (figure 6 - figure supplement 4), in part because there is little variability within countries and differences between countries are confounded by transmission intensity in this analysis. Surprisingly, there was no discernible association between either ITN retention times or the mean duration that nets remain in use on the impact of increasing campaign frequencies (figure 6 - figure supplements 6.F, 5.F). This is because ITNs have a bigger epidemiological impact in areas where they are used for longer, though policy-makers may wish to consider equity in disease burden in their tailoring process which is not considered here. There were, however, clearer relationships that indicated more cases may be averted following a switch from pyrethroid-PBO to pyrethroid-chlofenapyr triennial distribution in regions with greater mean access, use and use given access (figure 6 - figure supplements 8.H, 7.H, 9.H), though these associations are weaker than those seen for clinical cases or prevalence.

The cessation of mass campaigns is projected to lead to notable resurgences in clinical cases in many regions; an additional 150 to 300 clinical cases per 1,000 people are estimated for moderate to high transmission intensity regions (those with more than 150 mean annual clinical cases per 1,000 people under triennial pyrethroid-PBO distribution strategies). It has been suggested that the potential increase in cases caused by halting mass campaigns could be mitigated by concurrently switching to using pyrethroid-chlofenapyr ITNs in continuous distribution channels (figure 6.G vs. D). This modeling exercise indicates such mitigation is unlikely, with only 1 of the 146 regions showing a reduction of cases when switching from combined continuous distribution and triennial mass campaigns with pyrethroid-PBO to continuous distribution alone with pyrethroid-chlorfenapyr. This was for the capital city of Mali, Bamako (figure 5.C.vii), which has moderately low transmission, while Mali was estimated to distribute more ITNs through continuous channels than most other countries considered in this analysis (figure 2 - figure supplement 2).

Discussion

Earlier version: International guidelines on the use of ITNs have changed over the past two decades, moving from a universal target of 80% overall use in regions at risk of malaria to more nuanced guidance of “optimal” coverage in areas with resource limitations (World Health Organization, 2024b). This was an important policy change as it allowed NMPs to consider the use of alternative interventions which may be more cost-effective than pushing for very high population use in areas where this might be unachievable. Funding constraints have also increased the need for consideration of subnational tailoring, with many recommendations being made on the basis of transmission intensity in the World Health Organization (2025) Subnational Tailoring Reference Manual. However, a key uncertainty in assessing the potential impact of different ITN interventions has been how long nets remain in use rather than how long they are retained, and how this varies between regions. Here, to our best knowledge, we present the first estimates of subnational variation in ITN retention and the duration that ITNs remain in use, and also quantify for the first time how ITN use, access and retention vary between subnational regions across multiple African countries. Our work supports the change in guidance to optimal coverage as it highlights ITN interventions have notable differences in impact between settings, and that distributing fewer but more effective ITNs, particularly pyrethroid-chlorphenapyr products, is likely to be more impactful than maximising long-term coverage through increased campaign frequencies with pyrethroid-only ITNs. Our work also broadly supports World Health Organization (2025) recommendations for subnational tailoring, particularly the consideration of deprioritisation of ITN distribution in very low transmission settings. However, our results provide new indications that deprioritisation of areas with higher ITN use given access may lead to greater resurgences in cases, highlighting that subnational tailoring decisions could be optimised further by considering additional factors to transmission intensity alone. The modelling results indicate that the overall 80% use and access threshold was seldom consistently achieved across many countries in sub-Saharan Africa, and though this threshold appears relatively arbitrary, it is unclear how the new guideline of ‘optimal’ ITN use should be assessed. Mean use of ITNs under triennial mass campaigns with supplementary continuous distribution was estimated to be, on average, 41.7% (95% CrI: 39.6, 43.7) over all regions included in our analysis. This highlights that status quo approaches are typically yielding mean levels of ITN use that are roughly half of minimum universal coverage targets. Unless more nets are distributed or people keep their nets for significantly longer, universal coverage has potentially always been an unobtainable goal between mass campaigns.

Revised version: International guidelines on the use of ITNs have changed over the past two decades, moving from a universal target of 80% overall use in regions at risk of malaria to more nuanced guidance of “optimal” coverage in areas with resource limitations (World Health Organization, 2024b). This was an important policy change as it allowed NMPs to consider the use of alternative interventions which may be more cost-effective than pushing for very high population use in areas where this might be unachievable. Funding constraints have also increased the need for consideration of subnational tailoring, with many recommendations being made on the basis of transmission intensity in the World Health Organization (2025) Subnational Tailoring Reference Manual. However, a key uncertainty in assessing the potential impact of different ITN interventions has been how long nets remain in use rather than how long they are retained, and how this varies between regions. Here, to our best knowledge, we present the first estimates of subnational variation in ITN retention and the duration that ITNs remain in use, and also quantify for the first time how ITN use, access and retention vary between subnational regions across multiple African countries. Our work supports the change in guidance to optimal coverage as it highlights ITN interventions have notable differences in impact between settings, and that distributing fewer but more effective ITNs, particularly pyrethroid-chlorphenapyr products, is likely to be more impactful than maximising long-term coverage through increased campaign frequencies with pyrethroid-only ITNs. Our work also broadly supports World Health Organization (2025) recommendations for subnational tailoring, particularly the consideration of deprioritisation of ITN distribution in very low transmission settings. However, our results provide new indications that deprioritisation of areas with higher ITN use given access may lead to greater resurgences in cases, highlighting that subnational tailoring decisions could be optimised further by considering additional factors to transmission intensity alone. The modelling results indicate that the overall 80% use and access threshold was seldom consistently achieved across many countries in sub-Saharan Africa, and though this threshold appears relatively arbitrary, it is unclear how the new guideline of ‘optimal’ ITN use should be assessed. Mean use of ITNs under triennial mass campaigns with supplementary continuous distribution was estimated to be, on average, 45.4% (95% CrI: 45.3, 45.5) over all regions included in our analysis. This highlights that status quo approaches are typically yielding mean levels of ITN use that are roughly half of minimum universal coverage targets. Increasing mass campaign frequency to every two years is estimated to increase mean use by 8.50% (95% CrI: 8.47, 8.54), and is therefore unlikely to close this gap. Larger increases in mean use could plausibly be predicted under alternative ITN loss functions, such as the sigmoidal loss function used by Bhatt et al. (2015) and Bertozzi-Villa et al. (2021). However, inspection of ITN age distributions from DHS surveys provided limited evidence against the exponential decay function assumed here (appendix 2.2). Nevertheless, our estimates of mean use are consistent with previous modelling across a wider range of malaria-endemic countries by Bertozzi-Villa et al. (2021), where use was estimated to have broadly plateaued between 40% and 50% from approximately 2015 onwards, despite the use of a different loss function. Unless more nets are distributed or people keep their nets for significantly longer, universal coverage has potentially always been an unobtainable goal between mass campaigns.

The benefit of ITNs depends on how long people keep and use them. We therefore aimed to identify systematic differences in subnational use of and access to ITNs. Our central estimates indicated that 86.3% of subnational regions had a mean duration of access less than three years. Although the subnational regions included in this study are from a subset of countries in sub-Saharan Africa, our findings are broadly consistent with previous model-derived estimates of national median retention times, which have been estimated to be less than three years for 87% (Bertozzi-Villa et al., 2021) to 100% (Bhatt et al., 2015) of sub-Saharan African countries. Since we assume the duration of ITN use and access is exponentially distributed, our models incorporate a greater probability that a net will be lost, or cease to be used, in the immediate period after receipt. Therefore, we would argue our mean estimates are more comparable to the median estimates by Bertozzi-Villa et al. (2021) and Bhatt et al. (2015), which both assume a sigmoidal decay function. Although some field studies have reported ITN retention times greater than three years (Obi et al., 2020; Diouf et al., 2022), results differ by location, with the majority of field studies (for exmaple, Gnanguenon et al., 2014; Solomon et al., 2018; Lorenz et al., 2020) estimating that ITN retention fails to exceed three years in most settings. Even in isolation, our projections of additional cases averted under biennial over triennial campaigns raise further questions on whether the status quo of three-year intervals is sufficient to maintain adequate or “optimal” coverage; especially given its original basis stems from the duration of pyrethroid-class insecticidal activity estimated two decades ago (World Health Organization, 2005, 2013, 2023).

While previous work on net retention has largely focused on access, here we have also estimated the mean duration for which ITNs remain in use, which has greater epidemiological importance. Median estimates of the mean duration of use were less than three years in all but one of the subnational regions investigated. The work highlights that there is a faster rate of loss of use than access. Although the causes of this are unclear, it could be driven by poor quality ITNs, with people with older nets potentially not using them if degraded ITNs are perceived to provide lower protection. Indeed, reduced ITN integrity has been associated with reduced probability of use (Hiruy et al., 2021). ITN age is thought to be a better predictor of infection risk than durability alone (Andronescu et al., 2019). Therefore, even if duration of use is lower in some regions due to reduced ITN integrity, SBC campaigns either to promote use or to help with net care may still be beneficial.

We characterised historical trends in overall ITN use and access over time in different subnational regions. These estimated a gradual increase in both overall use and access over time, with waxing and waning dynamics driven by timings of mass campaigns; similar national patterns of declining use following mass campaigns have also been estimated in models by (Bhatt et al., 2015) and recorded in longitudinal field surveys (Hamre et al., 2020) and cluster randomised control trials (Protopopoff et al., 2018). Furthermore, surveillance data in multiple settings have also highlighted notable increases in malaria cases in the second year following mass campaigns (Zhou et al., 2016; Wotodjo et al., 2017; Girond et al., 2018; Topazian et al., 2021), providing further, albeit indirect, indications that overall ITN use declines notably in real-world settings by the second year following mass campaigns.

Although neither overall ITN use nor access in any subnational regions were projected to be maintained above the universal coverage threshold of 80%, mean use given access under present triennial strategies was estimated to be greater than 80% for almost half of subnational regions. This highlights that two decades after their widespread introduction, ITNs remain a valued commodity and key weapon in the expanding arsenal of malaria interventions when individuals have access to them. Moreover, if DHS surveys were primarily conducted in dry seasons in some locations (which is often the case), when ITN use given access is typically lower, our estimates in turn may be underestimated in some settings, especially in Senegal and regions of Mali, which are thought to have more seasonal fluctuations in use (Koenker et al., 2019b). However, to achieve consistent attainment of high population use targets, alternative distribution strategies, such as enhanced continuous distribution channels may be needed (Koenker et al., 2023b).

By identifying subnational regions with systematic low use given access and short durations of access, we highlight geographical areas where ITNs are likely to perform less effectively, irrespective of the local transmission intensity. This may aid NMPs in tailoring subnational strategies aside from those related to ITN distribution. For example, in regions where local knowledge indicates low use given access is driven by social and/or behavioural factors, SBC communication may be beneficial. Local understanding of the barriers to ITN use will be instrumental in assessing the cost-effectiveness of SBC strategies. While SBC communications, such as ‘hang-up’ campaigns have led to reported 250% relative increases in use in Zambia (Macintyre et al., 2012) and absolute increases between 11% and 17% in Togo (Desrochers et al., 2014) and Nigeria (Kilian et al., 2016), in other settings, such as in Uganda (Kilian et al., 2015), significant increases in use was not observed. We have not investigated the underlying factors that may have led to poor utilisation of ITNs in some subnational regions. We assume that use, access, and use given access immediately following a campaign, as well as their decay rates, remain unchanged when shifting from triennial to biennial strategies, and further work is needed to understand how use may change for different distribution strategies. Importantly, poor utilisation may arise due to ITN degradation (Mboma et al., 2021). An assessment of ITN durability studies across seven sub-Saharan African countries by Briet et al. (2020) found greater variability in durability between locations than between different net brands. Damaged ITNs may contribute not only to reduced utilisation, but also in reduced efficacy, particularly in areas with pyrethroid-resistant mosquitoes (Wheldrake et al., 2024).

Earlier version: Subnational tailoring of ITN interventions is increasingly recommended, with some World Health Organization (2024b,a, 2025) guidelines currently advising lower transmission settings to be considered first for deprioritisation from mass campaigns. Optimal subnational tailoring may not solely rely on deprioritisation, but may also require repriotisation of resources to other regions. Increasing the frequency of mass campaigns or procuring more effective, but likely more expensive, ITNs will inevitably increase the costs of conducting mass campaigns if the same level of coverage is achieved. If budget constraints necessitate the deprioritisation of campaigns, our results highlight that this should be avoided, if possible, in regions with moderate to high transmission intensity, particularly those with mean annual incidence exceeding 100–150 clinical cases per 1,000 people. Shortening campaign intervals from three to two years in moderate- and high-transmission regions is projected to avert more cases than the additional cases that may arise from ceasing campaigns in some lower-transmission settings. Additionally, although pyrethroid–chlorfenapyr ITNs are more costly, the additional cases projected to be averted by them relative to pyrethroid-only and pyrethroid–PBO ITNs are substantial. In certain national contexts it may be more cost-effective for biennial pyrethroid-chlorfenapyr campaigns to be conducted in fewer subnational regions even under reduced budgets. However, more thorough economic analyses will be needed to understand this fully. Moreover, as ITNs remain one of the most cost-effective malaria control interventions, improving the impact of them could still be more cost-effective than the introduction of new untested interventions (Topazian et al., 2023; Schmit et al., 2024).

Revised version: Subnational tailoring of ITN interventions is increasingly recommended, with some World Health Organization (2024b,a, 2025) guidelines currently advising lower transmission settings to be considered first for deprioritisation from mass campaigns. Optimal subnational tailoring may not solely rely on deprioritisation, but may also require repriotisation of resources to other regions. Increasing the frequency of mass campaigns or procuring more effective, but likely more expensive, ITNs will inevitably increase the costs of conducting mass campaigns if the same level of coverage is achieved. If budget constraints necessitate the deprioritisation of campaigns, our results highlight that this should be avoided, if possible, in regions with moderate to high transmission intensity, particularly those with mean annual incidence exceeding 100–150 clinical cases per 1,000 people. However, current incidence should be considered alongside the factors that gave rise to that transmission intensity, with caution exercised when deprioritising mass campaigns in areas where historically higher transmission may currently be suppressed by high ITN access, high use given access, or other interventions. Shortening campaign intervals from three to two years in regions with current moderateor high-transmission is projected to avert more cases than the additional cases that may arise from ceasing campaigns in some lower-transmission settings. Additionally, although pyrethroid–chlorfenapyr ITNs are more costly, the additional cases projected to be averted by them relative to pyrethroid-only and pyrethroid–PBO ITNs are substantial. In certain national contexts it may be more cost-effective for biennial pyrethroid-chlorfenapyr campaigns to be conducted in fewer subnational regions even under reduced budgets. However, more thorough economic analyses will be needed to understand this fully. Moreover, as ITNs remain one of the most cost-effective malaria control interventions, improving the impact of them could still be more costeffective than the introduction of new untested interventions (Topazian et al., 2023; Schmit et al., 2024).

In the context of finite procurement budgets, our framework identified subnational regions where increasing the frequency of mass campaigns or switching to more effective ITNs may have the greatest impact. Although we only included six countries in our analysis, they not only cover a wide geographical range, but also a broad spectrum of P. falciparum transmission intensities and estimated levels of pyrethroid resistance (Hancock et al., 2020). For all countries considered, transmission intensity was positively correlated with increased benefit of both switching to two-year mass campaign cycles or more efficacious ITN classes. The framework can be extended to other countries as DHS data becomes more available. In the absence of this, our results indicate areas with greater transmission intensity are likely to see the greatest impact from either increasing mass campaign frequencies or switching to pyrethroid-PBO or pyrethroid-chlorfenapyr ITNs, supporting WHO guidelines World Health Organization (2024a). Ultimately, decisions need to be made by NMPs whether they want to target the most cases, or whether decisions should be made with respect to equity of resources (does everyone have similar access to a net) or equity of impact (does everyone have a similar chance of getting malaria), both within and between regions.

As the toolbox for malaria control continues to expand, there is increasing need for ongoing monitoring and evaluation to facilitate evidence-based decision making and optimal deployment of interventions (Wilson et al., 2020). Longitudinal standardised surveys, such as the DHS, have the potential to be instrumental in assessing the effectiveness of interventions in real world settings, and have the potential to be used to refine distribution strategies to maximise impact. Such punctuated surveys can aid in monitoring the use of different public health products, such as spatial repellents or traps, and could independently assess the success of drug and vaccine campaigns. However, the timing of DHS surveys every few years creates challenges for assessing the effectiveness of both ITNs and other commodities. For example, if two DHS surveys were conducted at different times in relation to the last mass campaign, underlying trends in ITN use or access may fail to be captured. Similarly, if data are collected irregularly in relation to other future punctuated interventions, such as intermittent rapid diagnostic testing following mass vaccination campaigns, the methodology presented here provides a framework where potential biases can be controlled for, provided estimates of the duration of vaccine-induced immunity and the timings of vaccination campaigns are known. With appropriate adaptation to our model, this framework can facilitate more reliable inferences on the effectiveness of interventions beyond malaria than would be possible from survey data alone.

Earlier version: In conclusion, the work indicates that universal coverage targets of 80% are unlikely to be consistently met due to waning overall ITN use in the intervening years between triennial mass campaigns. Improved coverage can be achieved through more frequent biennial distributions, though this is unlikely to be feasible at scale given the current funding landscape. Indeed, when resources are constrained, deprioritisation of ITN mass campaigns in certain settings is being increasingly considered through subnational tailoring of malaria control interventions. Our work highlights that the relationship between transmission intensity (whether measured in terms of prevalence or clinical cases) and intervention impact is non-linear, and notable resurgences in cases may follow when campaigns are deprioritised in all but very low transmission settings. This broadly supports WHO subnational tailoring guidance, which suggests consideration of deprioritising distribution of ITNs in regions with Pf PR2−10 < 1% (World Health Organization, 2025). However, while the World Health Organization (2025) Subnational Tailoring Reference Manual proposes that the withdrawal of ITNs in favour of indoor residual spraying should be considered in areas with low ITN use, here we estimate that ITN use alone appears to be a notably poorer predictor of the impact of ceasing mass campaigns than use given access. Our findings suggest that regions with higher use given access may experience disproportionately greater resurgences in cases following deprioritisation. This implies that regions with low use given access may warrant consideration for cessation of ITN distribution, rather than decisions being based solely on low overall ITN use irrespective of whether communities have sufficient ITN access. However, subnational differences in ITN use, access and retention are key knowledge gaps in many settings, and when estimated from infrequent surveys they are highly sensitive to bias arising from the timing of surveys relative to when campaigns were conducted. To our knowledge, this study is the first to estimate subnational variation in ITN retention and the first to estimate the duration that ITNs remain in use, which is of greater epidemiological relevance than retention time. It also provides a novel framework to correct for biases in estimates of ITN use and access arising from when campaigns were conducted. Although campaigns have historically aided increasing ITN use and access over time, we estimate the mean duration of ITN use is consistently shorter than mean retention times in all regions. This raises questions about whether punctuated distribution of ITNs through campaigns is the optimal mechanism for maximising their effectiveness and cost-effectiveness. Maximising the cost-effectiveness of interventions has become increasingly pertinent in the current funding context, and consideration of alternative distribution strategies, such as increased distribution through continuous distribution channels, including school- or community-based distribution, may be warranted. Frameworks such as the one presented here, which take into account the potential for impact from different net types and the high variability of ITN duration and use, could support NMP decision-making on how best to maximise impact from available funds. Whilst such frameworks may be a useful tool, local knowledge of factors impacting ITN access and use as well as operational decision making will be paramount for NMP-led tailoring of subnational strategies.

Revised version: In conclusion, the work indicates that universal coverage targets of 80% are unlikely to be consistently met due to waning overall ITN use in the intervening years between triennial mass campaigns. Improved coverage can be achieved through more frequent biennial distributions, though this is unlikely to be feasible at scale given the current funding landscape. Indeed, when resources are constrained, deprioritisation of ITN mass campaigns in certain settings is being increasingly considered through subnational tailoring of malaria control interventions. Our work highlights that the relationship between transmission intensity (whether measured in terms of prevalence or clinical cases) and intervention impact is non-linear, and notable resurgences in cases may follow when campaigns are deprioritised in all but very low transmission settings. This broadly supports WHO subnational tailoring guidance, which suggests consideration of deprioritising distribution of ITNs in regions with Pf PR2−10 < 1% (World Health Organization, 2025). However, while the World Health Organization (2025) Subnational Tailoring Reference Manual proposes that the withdrawal of ITNs in favour of indoor residual spraying should be considered in areas with low ITN use, here we estimate that ITN use alone appears to be a notably poorer predictor of the impact of ceasing mass campaigns than use given access. Our findings suggest that regions with higher use given access may experience disproportionately greater resurgences in cases following deprioritisation. This implies that regions with low use given access may warrant consideration for cessation of ITN distribution, rather than decisions being based solely on low overall ITN use irrespective of whether communities have sufficient ITN access. However, subnational differences in ITN use, access and retention are key knowledge gaps in many settings, and when estimated from infrequent surveys they are highly sensitive to bias arising from the timing of surveys relative to when campaigns were conducted. To our knowledge, this study is the first to estimate subnational variation in ITN retention and the first to estimate the duration that ITNs remain in use, which is of greater epidemiological relevance than retention time. It also provides a novel framework to correct for biases in estimates of ITN use and access arising from when campaigns were conducted. Although campaigns have historically aided increasing ITN use and access over time, we estimate the mean duration of ITN use is consistently shorter than mean retention times in all regions. This raises questions about whether punctuated distribution of ITNs through campaigns is the optimal mechanism for maximising their effectiveness and cost-effectiveness. Maximising the cost-effectiveness of interventions has become increasingly pertinent in the current funding context, and consideration of alternative distribution strategies, such as increased distribution through continuous distribution channels, including schoolor community-based distribution, may be warranted. Frameworks such as the one presented here, which take into account the potential for impact from different net types and the high variability of ITN duration and use, could support NMP decision-making on how best to maximise impact from available funds. Whilst such frameworks may be a useful tool, local knowledge of factors impacting ITN access and use as well as operational decision making will be paramount for NMP-led tailoring of subnational strategies.

Methods and Materials

Inclusion criteria and data sources

The majority of data were sourced from Demographic and Health Surveys (DHS) and Malaria Indicator Surveys (MIS) (ICF, 2025). AIDS Indicator Surveys (AIDS) were also included where additional malaria questions were included (ICF, 2025). Subsequent references to DHS surveys refer to all surveys included. Data on the number of ITNs distributed nationally were also sourced from (The Alliance for Malaria Prevention, 2024).

Here we assume each ITN is used every night for a continuous period after receipt (duration of use), but may remain accessible for a different length of time following receipt (duration of access/retention time). It is important to distinguish between people who have access to an ITN but don’t sleep beneath it and those who don’t have access to one. We generate subnational estimates of both the population mean duration of use and retention time and assume these are constant over time. Single estimates of both these quantities are estimated separately for rural and urban areas of each region to allow the identification of systematic differences between locations that could allow the geographical targeting of resources.

Retrospective analyses were conducted from 2008, to align with the introduction of mass distribution campaigns, until 2024. Analyses were conducted at an administrative-one level stratified by urban and rural settings, as recorded in survey responses. To account for inconsistencies in administrative-one naming conventions, such as alternative French and English spelling, best efforts were made to unify recorded administrative-one entries with the naming conventions recorded by GADM v4.0 (Global Administrative Areas, 2021). Note that while administrative-one is referred to as a region in Senegal, for example, terminology differs between countries; here, the units of analysis are consistently referred to as (subnational) regions. The initial inclusion criteria were restricted to countries where at least three DHS/MIS surveys had been conducted between 2010 and 2022. This was to ensure there were sufficient data for all countries after mass campaigns had been established. Ten countries initially met this criteria: Burkina Faso, Ghana, Liberia, Malawi, Mali, Mozambique, Nigeria, Senegal, Tanzania and Uganda.

All mass campaigns were assumed to be conducted nationally, although their timings were not restricted to occur in synchrony across subnational regions. Instead, the timing of each mass campaign was allowed to vary subnationally within time intervals that were defined at a national level. Due to this assumption, Liberia was excluded from our analysis due to rolling campaigns being conducted between 2008 and 2012 (Koenker et al., 2019a). We also excluded Nigeria where mass campaigns follow different distribution schedules between states (Koenker et al., 2019a). Tanzania was also excluded as a mixture of annual school-based distributions and five-yearly mass campaigns are conducted within different subnational regions (Koenker et al., 2023a). We finally excluded Uganda since there have been notable changes to its administrative-one level boundaries over the assessed period; this has resulted in significant discrepancies between the administrative-one level regions recorded in DHS surveys and the naming conventions used in GADM v4.0. Although Ghana was included in all analyses, the sub-division of the regions of Brong-Ahafo, Northern and Volta in 2019 was not accounted for in the retrospective analyses of historical use and access; the new regions were however accounted for in the transmission dynamics models for future projections. For Malawi, we conducted retrospective analyses of use and access at an administrative-one level (termed regions in the country); however, as data we sourced for other covariates, such as rainfall patterns and non-ITN interventions, were stratified at an administrative-two level (districts) for Malawi (Winskill, 2024b), we excluded the country from the transmission dynamics analyses. These other covariates were stratified at the administrative-one level for all other countries (Winskill, 2024b).

All subsequent analyses described are conducted independently for use and access and were conducted in R (v4.3.2, R Core Team, 2023). The code for models used in this study can be found at https://github.com/andrewcglover/itn_campaigns. Unless otherwise stated, median central estimates are quoted in the main text.

Historical use, access and retention times

Earlier version: Due to the sparsity and irregularity of DHS and MIS surveys, we were unable to investigate seasonal fluctuations in either access or use; we therefore assume that nets provide access or are used continuously over some period of time. We do not assume these periods are equal, and so differentiate between the duration of time that a net is slept under (duration of use) and the period for which it provides access before being discarded or repurposed (retention time). Using a Bayesian hierarchical model fitted to pooled ITN age distribution data across DHS surveys, we separately generated initial subnational estimates of the mean duration of use and mean retention times of ITNs from DHS survey data using the age distributions of ITNs in surveyed households. When inferring mean retention times under our hierarchical model, we implicitly assume ITNs are only distributed to individuals without them; this is a necessary methodological step, but is relaxed later in our analysis. Although the methodology is applied analogously in the context of use, the remainder of this section focuses on access.

Revised version: Due to the sparsity and irregularity of DHS and MIS surveys, we were unable to investigate seasonal fluctuations in either access or use; we therefore assume that nets provide access or are used continuously over some period of time. Although seasonal differences are not identifiable under our model, overdispersion in individual-level use and access may partially absorb unexplained variation between observations, including variation arising from unmodelled seasonal differences (equation 7). However, this does not allow us to quantify seasonal variation in ITN use or access. If ITN use were systematically higher during high-transmission rainy seasons, our assumption of continuous use may underestimate the protective impact of ITNs during these periods. We do not assume these periods are equal, and so differentiate between the duration of time that a net is slept under (duration of use) and the period for which it provides access before being discarded or repurposed (retention time). Using a Bayesian hierarchical model fitted to pooled ITN age distribution data across DHS surveys, we separately generated initial subnational estimates of the mean duration of use and mean retention times of ITNs from DHS survey data using the age distributions of ITNs in surveyed households. When inferring mean retention times under our hierarchical model, we implicitly assume ITNs are only distributed to individuals without them; this is a necessary methodological step, but is relaxed later in our analysis. Although the methodology is applied analogously in the context of use, the remainder of this section focuses on access.

Earlier version: Under the assumption that nets are lost at a constant rate, the long-term age distribution of nets providing access follows an exponential distribution, allowing retention times to be inferred from observed net ages. The hierarchical structure pools information across regional, national, and continental spatial scales to reduce bias arising from survey timing relative to mass campaigns, particularly in regions with sparse survey data. Region-specific mean retention times were modelled with country-level random effects, and country-level means were in turn modelled with a continent-level random effect. Right-censoring of reported net ages at 36 months and survey sampling weights were explicitly accounted for in the likelihood when fitting the model in RStan (v2.32.6, Stan Development Team (2024)). Full specification of the hierarchical model and prior distributions are provided in appendix 2.

Revised version: Under the assumption that nets are lost at a constant rate, the long-term age distribution of nets providing access follows an exponential distribution, allowing retention times to be inferred from observed net ages. The hierarchical structure pools information across regional, national, and continental spatial scales to reduce bias arising from survey timing relative to mass campaigns, particularly in regions with sparse survey data. Region-specific mean retention times were modelled with country-level random effects, and country-level means were in turn modelled with a continent-level random effect. Right-censoring of reported net ages at 36 months and survey sampling weights were explicitly accounted for in the likelihood when fitting the model in Stan via the R interface RStan (v2.32.6, Stan Development Team (2024)). Full specification of the hierarchical model and prior distributions are provided in appendix 2.

Earlier version: We then sought to estimate the timings of mass campaigns in different regions. We use annual estimates of the number of ITNs delivered nationally from the AMP Net Mapping Project (The Alliance for Malaria Prevention, 2024) to estimate time intervals within which each mass campaign was likely to occur for each country. We do however allow for the timings of campaigns to vary sub-nationally within these intervals. For each region, the timing of each mass campaign is estimated probabilistically by constructing distributions over plausible months using DHS data on the ages of ITNs recorded in surveys. As time elapses since a mass campaign, ITNs are less likely to be recorded in DHS surveys. Our initial ITN retention estimates help correct for this bias (see appendix 3.1 for further details). We can then probabilistically infer the timing of the previous campaign (ϕj) relative to some time tj in a given region (see appendix 3.2 for further details), noting that in contrast to appendix 3.2, we do not explicitly indicate subscript i notation to indicate the region of interest for notational ease. In turn, this allows us to infer the time since the previous campaign (mj) relative to all calendar times considered in our model: (1) mj=tj−ϕj We then utilise the following discrete-time model of the proportion of the population with access (pj) for a given region at some time tj in the observed survey period: (2) pj=cj+dj (3) cj=q0pj0exp⁡(−λmj) (4) dj=(1−q0)pj0 (5) pj0=ψlogit−1(β0+βttj) where superscripts 0 and t are used for notational purposes and do not indicate exponentiation. An in-depth explanation of the model can be found in appendix 4.1, but in brief, we firstly assume access can be attributable to ITNs received from campaign (cj) or continuous distribution (dj) channels (equation (2)). Secondly, we use q0 to denote the proportion of the population with access to an ITN from a campaign relative to access to any ITN immediately post-campaign, which is assumed to be the same following all campaigns for a given region (equation (3)); the relative proportion with access to an ITN from continuous distribution channels at these times is therefore given by 1 − q0 (equation (4)). Thirdly, we assume logistic-type growth in access achieved immediately following each subsequent campaign, denoted by pj0, this growth is characterised by the parameters βj0 and βjt (inverse-logit function in equation (5)), and is bounded by an upper limit, ψ. We therefore also implicitly assume logistic-type growth in access to ITNs from continuous channels over time (pj0 in equation (4)) and ITNs distributed through campaigns over time (pj0 in equation (3)). Finally, we assume access to campaign nets decays following mass campaigns (exponential term in equation (3)); the rate of this decay is governed by the parameter λi which is related to the mean retention time (Λ−1) such that: (6) Λ≈q0λ where retention times are assumed to be the same for all nets irrespective of their distribution channel. In contrast to the earlier hierarchical model, the discrete-time model and subsequent intervention impact projections assumes random allocation of ITNs in deriving the updated estimates of mean retention time (equation (6)). Under this assumption, at the point of distribution, an individual already with access to an ITN has the same probability of receiving a new ITN as an individual currenlty without access. Further details on how retention times are inferred from the approximation in equation (6) can be found in appendix 4.4.

Revised version: We then sought to estimate the timings of mass campaigns in different regions. We use annual estimates of the number of ITNs delivered nationally from the AMP Net Mapping Project (The Alliance for Malaria Prevention, 2024) to estimate time intervals within which each mass campaign was likely to occur for each country. We do however allow for the timings of campaigns to vary subnationally within these intervals. For each region, the timing of each mass campaign is estimated probabilistically by constructing distributions over plausible months using DHS data on the ages of ITNs recorded in surveys. As time elapses since a mass campaign, ITNs are less likely to be recorded in DHS surveys. Our initial ITN retention estimates help correct for this bias (see appendix 3.1 for further details). We can then probabilistically infer the timing of the previous campaign (ϕj) relative to some time tj in a given region (see appendix 3.2 for further details), noting that in contrast to appendix 3.2, we do not explicitly indicate subscript i notation to indicate the region of interest for notational ease. In turn, this allows us to infer the time since the previous campaign (mj) relative to all calendar times considered in our model:

Added in the revised version: (1) mj=tj−ϕj

Added in the revised version: We then utilise the following discrete-time model of the proportion of the population with access (pj) for a given region at some time tj in the observed survey period:

Added in the revised version: (2) pj=cj+dj

Added in the revised version: (3) cj=q0pj0exp⁡(−λmj)

Added in the revised version: (4) dj=(1−q0)pj0

Added in the revised version: (5) pj0=ψlogit−1(β0+βttj)

Added in the revised version: where superscripts 0 and t are used for notational purposes and do not indicate exponentiation. An in-depth explanation of the model can be found in appendix 4.1, but in brief, we firstly assume access can be attributable to ITNs received from campaign (cj) or continuous distribution (dj) channels (equation (2)). Secondly, we use q0 to denote the proportion of the population with access to an ITN from a campaign relative to access to any ITN immediately post-campaign, which is assumed to be the same following all campaigns for a given region (equation (3)); the relative proportion with access to an ITN from continuous distribution channels at these times is therefore given by 1 − q0 (equation (4)). Thirdly, we assume logistic-type growth in access achieved immediately following each subsequent campaign, denoted by pj0, this growth is characterised by the parameters βj0 and βjt (inverse-logit function in equation (5)), and is bounded by an upper limit, ψ. We therefore also implicitly assume logistic-type growth in access to ITNs from continuous channels over time (pj0 in equation (4)) and ITNs distributed through campaigns over time (pj0 in equation (3)). Finally, we assume access to campaign nets decays following mass campaigns (exponential term in equation (3)); the rate of this decay is governed by the decay parameter, λ, which is related to the mean retention time (Λ−1) such that:

Added in the revised version: (6) Λ≈q0λ

Added in the revised version: where retention times are assumed to be the same for all nets irrespective of their distribution channel. In contrast to the earlier hierarchical model, the discrete-time model and subsequent intervention impact projections assumes random allocation of ITNs in deriving the updated estimates of mean retention time (equation (6)). Under this assumption, at the point of distribution, an individual already with access to an ITN has the same probability of receiving a new ITN as an individual currenlty without access. Further details on how retention times are inferred from the approximation in equation (6) can be found in appendix 4.4.

We do however use our regional posterior estimates of mean ITN retention times from our previous hierarchical model (which assumed ITNs are only distributed to those without access) as informative priors for mean retention times under random allocation to aid parameter identifiability. Either non-informative or regularising priors are used for other region-specific parameters (see appendix 4.2 for further details).

Earlier version: When fitting our model in RStan (v2.32.6, Stan Development Team (2024)), we accounted for the potential of individual-level access (and indeed use, in that context) to vary subregionally; for example, between different wards or districts. Indeed, examination of the raw data (for example, figure 2 - figure supplement 1.A) indicates substantial variability in the recorded proportions of use and access in adjacent months. This may reflect substantial subregional variability, and there are likely other covariates that our model does not account for that will create additional variability in the probability of access (and use) within an administrative-one unit at any given point in time. Our discrete-time model thus far describes the proportion of the population with access (pj in equation (2)). However, we also consider the probability of individuals having access at each time tj, which is denoted by p~j. We do not assume the probability of access is the same for all individuals in a region at a given point in time. Instead, we assume the probability any given individual has access to an ITN at time tj can be described by a Beta distribution: (7) p~j∼Beta(α0pj,α0(1−pj)) Under this assumption, the mean probability of access across all individuals at a given time is equal to the overall proportion of the population with access. Meanwhile, for a given proportion of the population with access, the overdispersion parameter, denoted by α0 is approximately inversely proportional to the variance in the probability of individual-level access. This overdispersion parameter, which is fitted for each region and assumed to remain constant for each over time, therefore controls the estimated level of subregional heterogeneity in access. DHS data on the total number of surveyed individuals by month, in addition to the number with access to ITNs from different sources, are accounted for in our likelihood specification for region-specific access over time and the relative contribution towards this from campaign and continuous distribution channels. Further details can be found in appendix 4.3.

Revised version: When fitting our model in Stan with the R interface RStan (v2.32.6, Stan Development Team (2024)), we accounted for the potential of individual-level access (and indeed use, in that context) to vary subregionally; for example, between different wards or districts. Indeed, examination of the raw data (for example, figure 2 figure supplement 1.A) indicates substantial variability in the recorded proportions of use and access in adjacent months. This may reflect substantial subregional variability, and there are likely other covariates that our model does not account for that will create additional variability in the probability of access (and use) within an administrative-one unit at any given point in time. Our discrete-time model thus far describes the proportion of the population with access (pj in equation (2)). However, we also consider the probability of individuals having access at each time tj, which is denoted by p~j. We do not assume the probability of access is the same for all individuals in a region at a given point in time. Instead, we assume the probability any given individual has access to an ITN at time tj can be described by a Beta distribution:

Added in the revised version: (7) p~j∼Beta(α0pj,α0(1−pj))

Added in the revised version: Under this assumption, the mean probability of access across all individuals at a given time is equal to the overall proportion of the population with access. Meanwhile, for a given proportion of the population with access, the overdispersion parameter, denoted by α0, is approximately inversely proportional to the variance in the probability of individual-level access. This overdispersion parameter, which is fitted for each region and assumed to remain constant for each over time, therefore controls the estimated level of subregional heterogeneity in access. DHS data on the total number of surveyed individuals by month, in addition to the number with access to ITNs from different sources, are accounted for in our likelihood specification for region-specific access over time and the relative contribution towards this from campaign and continuous distribution channels. Further details can be found in appendix 4.3.

With the exception of panels B and C respectively in figures 2 and 5 (and associated supplementary figures), figures 1 - 5 are generated using predictive posterior estimates from our discrete time model. Current mean estimates under triennial (or biennial) campaigns (figure 3) are inferred from equations (2) - (5) over a three- (or two-) year period had a hypothetical campaign been conducted in every region at the start of January-2025, using posterior parameter estimates for December-2024 (the end of the period of included surveys). Estimates of the proportion of the population with different probabilities of access (and use) are then inferred from this using equation (7) (figure 4). This provides a measure of equity in use and access of nets within a region which will have important policy implications for distribution strategies.

Simulating future distribution strategies

Earlier version: The future scenarios considered were no future ITN distribution, continuous distribution in isolation, in addition to biennial and triennial campaigns with supplementary continuous distribution. Distribution of pyrethroid-PBO ITNs is assumed in all regions prior to December-2024. Distribution of different ITN classes (pyrethroid-only, pyrethroid-PBO or pyrethroid-chlorfenapyr) is simulated from January-2025 onwards for each distribution strategy. Biennial and triennial distribution strategies are respectively simulated to commence either 2- or 3-years following the median estimate of the last estimated campaign prior to 2025. For strategies with campaigns, clinical cases are recorded over a 6-year interval commencing on the first simulated campaign after December-2024. Clinical case estimates for continuous-only distribution and no future distribution are recorded over a 6-year period commencing on the mid-point between the start of the recorded 6-year biennial and triennial scenarios, should these differ for a region. For each subnational region, 100 stochastic realisations were simulated for a population of 100,000 individuals, and cases averted were calculated relative to the baseline scenario with no future ITN distribution. For each realisation, parameters are drawn from our joint posterior distribution of fitted parameters from our discrete-time model of use, however ongoing logisitic growth in use of ITNs from continuous channels and immediately following campaigns is not assumed beyond December-2024. Future campaigns, whether conducted every two or three years, are therefore assumed to achieve a consistent initial level of use. Since epidemiological dynamics are only dependent on whether ITNs are ultimately used, our simulations are agnostic of our access estimates. For each realisation and distribution event, parameter draws are also made from previously generated posterior distributions of the probability of repellency and mortality at different pyrethroid-resistance levels; these were fitted to data from systematic reviews of experimental hut trials across Africa (Sherrard-Smith et al., 2022; Churcher et al., 2024). Region-specific uncertainty in ITN efficacy, use, retention, and the relative contributions of continuous and campaign channels is therefore propagated through to our estimates of cases averted.

Revised version: The future scenarios considered were no future ITN distribution, continuous distribution in isolation, in addition to biennial and triennial campaigns with supplementary continuous distribution. Distribution of pyrethroid-PBO ITNs is assumed in all regions prior to December-2024. Distribution of different ITN classes (pyrethroid-only, pyrethroid-PBO or pyrethroid-chlorfenapyr) is simulated from January-2025 onwards for each distribution strategy. Biennial and triennial distribution strategies are respectively simulated to commence either 2or 3-years following the median estimate of the last estimated campaign prior to 2025. For strategies with campaigns, clinical cases are recorded over a 6-year interval commencing on the first simulated campaign after December-2024. Clinical case estimates for continuous-only distribution and no future distribution are recorded over a 6-year period commencing on the mid-point between the start of the recorded 6-year biennial and triennial scenarios, should these differ for a region. For each subnational region, 100 stochastic realisations were simulated for a population of 100,000 individuals, and cases averted were calculated relative to the baseline scenario with no future ITN distribution. For each realisation, parameters are drawn from our joint posterior distribution of fitted parameters from our discrete-time model of use, however ongoing logistic growth in use of ITNs from continuous channels and immediately following campaigns is not assumed beyond December-2024. Future campaigns, whether conducted every two or three years, are therefore assumed to achieve a consistent initial level of use. Since epidemiological dynamics are only dependent on whether ITNs are ultimately used, our simulations are agnostic of our access estimates. For each realisation and distribution event, parameter draws are also made from previously generated posterior distributions of the probability of repellency and mortality at different pyrethroid-resistance levels; these were fitted to data from systematic reviews of experimental hut trials across Africa (Sherrard-Smith et al., 2022; Churcher et al., 2024). Region-specific uncertainty in ITN efficacy, use, retention, and the relative contributions of continuous and campaign channels is therefore propagated through to our estimates of cases averted.

Regional differences in transmission intensities are accounted by calibration of historical baseline estimates of the entomological inoculation rate (EIR) prior to the introduction of any interventions. Baseline EIR estimates were calibrated by fitting our modelled annual mean Plasmodium falciparum prevalence in children between the ages of 2 and 10 (Pf PR2−10) to those estimated by the Malaria Atlas Project (2024) (figure 2.B). Forecasted increases in pyrethroid resistance were incorporated, while coverage of other interventions (including IRS, SMC, drug treatment and vaccination) was assumed to be maintained at 2024 levels in each region; these were estimated using the site (v0.2.2) R package (Winskill, 2024b) package

Future projections were generated using malariasimulation (v1.6.0) (Charles et al., 2024), an individual-based malaria transmission model that incorporates typical seasonal patterns in transmission. The transmission dynamics model, which incorporates typical seasonal patterns in each region, is broadly able to reproduce the observed changes in Pf PR6−59mo, as estimated by DHS surveys, providing some confidence in the resulting epidemiological predictions. However, while the model captures typical seasonal dynamics, it does not account for inter-annual variation or specific weather events. These factors may have contributed to larger deviations from DHS Pf PR6−59mo point estimates at certain times; estimates that themselves often have wide 95% credible intervals (figure 2.B). Full details of model calibration, ITN efficacy parameters, and simulation procedures are provided in appendix 5.